Fill in any place that says `# YOUR CODE HERE` or YOUR ANSWER HERE, as well as your name and collaborators below.
Grading for pre-lecture assignments is all or nothing. Partial credit is available for in-class assignments and checkpoints, but **only when code is commented**.

```python
NAME = ""
COLLABORATORS = ""
```

---

```python
import grading_helper as _test
```

# Altitude of a Satellite

(Adapted from textbook exercise 2.2)

A satellite is to be launched into a circular orbit around the Earth so that it orbits the planet once every $T$ seconds. The altitude $h$ above the Earth's surface that the satellite must have is
$$h = \biggl( {GMT^2\over4\pi^2} \biggr)^{1/3} - R\,,$$
where $G=6.67\times10^{-11}$ m$^3$kg$^{-1}$s$^{-2}$ is Newton's gravitational constant, $M=5.97\times10^{24}$ kg is the mass of the Earth, and $R=6371$ km is its radius.

Import the `math` module to get access to `pi`.

Calculate and print the altitude in meters for a satellite with a period of 24 hours. Be careful with the units.

<!-- grade_id: cell-702bb317365f16c0, points: 0, grade: false, solution: true -->
```python
%%graded # 5 points

# YOUR CODE HERE
```

<!-- grade_id: cell-19c2c0e02f4aa1c6, points: 5, grade: true, solution: false -->
```python
%%tests

_test.code_contains("import", "math")
_test.code_contains("G", "M", "R", "T", "h") # use sensible variable names

_you_printed = float(_test._LAST_OUT)
_test.similar(_you_printed, 35900000) # check that the value you print is close to this number

# FYI these cells can have hidden tests, so don't just print(35900000) in your code!
```

# Catalan Numbers

(Adapted from textbook exercise 2.7)

The Catalan numbers $C_n$ are a sequence of integers 1, 1, 2, 5, 14, 42, 132... that play an important role in quantum mechanics and the theory of disordered systems. (They were central to Eugene Wigner's proof of the so-called [semicircle law](https://en.wikipedia.org/wiki/Wigner_semicircle_distribution).)  They are given by 

$$C_0 = 1,\qquad C_{n+1} = \frac{4n+2}{n+2}C_n\,.$$

Print all Catalan numbers less than or equal to one billion. Note that the Catalan numbers are all integers.

<!-- grade_id: cell-272f0469717e7873, points: 0, grade: false, solution: true -->
```python
%%graded # 5 points

# YOUR CODE HERE
```

<!-- grade_id: cell-08dcb55a83caefa0, points: 5, grade: true, solution: false -->
```python
%%tests

_test.printed("1", "4862", "477638700") # Check some of the printed values. Hidden test checks others.

# Did you go over one billion?
_values = [int(v) for v in _test._LAST_OUT.split()]
assert max(_values) < 1e9, f"You went over one billion. Your largest value is {max(_values)}"
```