Grading for pre-lecture questions is all or nothing. Partial credit is available for in-class assignments and checkpoints. Partial credit will primarily be determined by the quality of the comments you make in your code. Tell me what you intend the code to do. If there are no comments, there is no partial credit!

```python
NAME = ""
COLLABORATORS = ""
```

---

# Learning Objectives

This lecture will show you how to:
1. Use arrays in your calculations
2. Identify the significance of an array's type
3. Quickly create large arrays
4. Index and slice arrays
5. Load data from text files

```python
# imports
import numpy as np # provides arrays for fast calculations

import grading_helper as _test
```

# Introduction to NumPy

```python
%video 4k9gFlkBeyU
```

Summary:
- Typical import statement: `import numpy as np`
- Create an array: `A = np.array([1, 2, 3, 4, 5])`
- Operations on arrays are element-by-element.
- Arrays have a few built-in methods that are useful for stats: `A.mean()` and `A.std()`

```python

```

## Array Example

A ball undergoing free-fall has $y$-components given by the equation
$$y = y_0 + v_0t + \frac12 at^2\,.$$
Let $y_0 = 0$, $v_0=10$, and $a=-9.8$ (in MKS units). Find the $y$-component of the ball as a function of time.

```python
%video cPbPxaxROB0
```

```python

```

## Your Turn

Given an array `x = np.array([1, 2, 3, 4, 5, 6, 7, 8, 9, 10])`, calculate an array `v` using the equation
$$v = \sqrt{v_0^2 + 2ax}\,,$$
where $v_0 = 5$ and $a=2$.

> Hint: you can calculate the square root of an array using `np.sqrt` (like any array operation, it will apply element-by-element).

```python
x = np.array([1, 2, 3, 4, 5, 6, 7, 8, 9, 10])
```

<!-- grade_id: cell-bc32951322025176, points: 0, grade: false, solution: true -->
```python
%%graded # 2 points

# YOUR CODE HERE
```

<!-- grade_id: cell-9c67503909aa3da0, points: 2, grade: true, solution: false -->
```python
%%tests

_test.equal(len(x), len(v))
_test.similar(v[0], 5.39, rtol=0.01)
_test.similar(v[-1], 8.06, rtol=0.01)
```

# Arrays Have a Type

```python
%video EgEwJA4WZhA
```

Summary:
- Every element of an array must have the same type (stored as the array's `dtype`).
- The `dtype` is inferred when the array is created, or can be explicitly set `A = np.array([1, 2, 3], dtype=float)`
- When different types of arrays are combined, the result is an array of the more general type.

```python

```

# Ways to Create Arrays

```python
%video nAyAFZ1TYdA
```

Summary:
- `np.arange` acts like `range`, and creates arrays with a given spacing.
- `np.linspace` is similar, but instead creates arrays with a certain number of equally-spaced elements.
- `np.zeros` creates an array of all zeros.
- `np.ones` creates an array of all ones.
- 2-D arrays can be created with `np.array([[1, 2, 3], [4,5,6], [7,8,9]])` or e.g. `np.zeros((3,3))`

```python

```

## Your Turn

Create an array named `theta` that contains 20 equally-spaced values ranging from 0 to $2\pi$, **including** both endpoints.

<!-- grade_id: cell-4ca510ccb139b2d5, points: 0, grade: false, solution: true -->
```python
%%graded # 1 point

# YOUR CODE HERE
```

<!-- grade_id: cell-852e5ffe89435181, points: 1, grade: true, solution: false -->
```python
%%tests

_test.equal(len(theta), 20)
_test.similar(theta[0], 0)
_test.similar(theta[-1], 2*np.pi)
_test.similar(theta[1], 2*np.pi/19)
```

# Array Indexing and Slicing

```python
%video PXERB4XfKXw
```

Summary:
- Indexing and slicing for 1-D arrays is the same as it was for strings, lists, and tuples.
- Indexing for 2-D arrays looks like `my_array[row, col]`. Arrays can also be larger than 2-D. Just add more indices.
- `:` means "everything", so `my_array[:, 0]` is everything in the first column.
- Arrays can also be sliced using comparisons. `A[A>0]` returns only the elements of `A` that are positive.

```python

```

## Your Turn

Use slicing to change
```
A = np.array([[1, 1, 1, 1, 1],
              [1, 1, 1, 1, 1],
              [1, 1, 1, 1, 1],
              [1, 1, 1, 1, 1],
              [1, 1, 1, 1, 1]])
```

into

```
A = np.array([[0, 0, 0, 0, 0],
              [0, 1, 1, 1, 0],
              [0, 1, 1, 1, 0],
              [0, 1, 1, 1, 0],
              [0, 0, 0, 0, 0]])
```
It should take you no more than four lines of code.

```python
A = np.array([[1, 1, 1, 1, 1],
              [1, 1, 1, 1, 1],
              [1, 1, 1, 1, 1],
              [1, 1, 1, 1, 1],
              [1, 1, 1, 1, 1]])
```

<!-- grade_id: cell-003cd58a38df7ef5, points: 0, grade: false, solution: true -->
```python
%%graded # 2 points

# YOUR CODE HERE
```

<!-- grade_id: cell-1ab4e5ea2e5cd57c, points: 2, grade: true, solution: false -->
```python
%%tests

_test.equal(A, [[0,0,0,0,0], [0,1,1,1,0], [0,1,1,1,0], [0,1,1,1,0], [0,0,0,0,0]])
_test.code_contains("array", forbidden=True) # makes sure you don't just recreate the array from scratch
assert type(A) == np.ndarray, "A should still be an array!"
```

# Loading Data From a Text File

```python
%video 8iwpsRKiwN0
```

Summary:
- `np.loadtxt` (simple) and `np.genfromtxt` (more options) create arrays from text files.
- Use `unpack=True` and tuple unpacking to put each column into its own 1-D array.

```python

```

# Additional Resources

- The official  NumPy documentation https://docs.scipy.org/doc/numpy/user/index.html
- Textbook section 2.4

The textbook has a bad habit of using for loops with lists instead of numpy arrays. For example, the author writes a lot of code that looks like
```
xpoints = []
ypoints = []
for x in linspace(0,10,100):
      xpoints.append(x)
      ypoints.append(sin(x))
```
That's not how it should be done. Instead, you should write
```
x = linspace(0,10,100)
y = sin(x)
```
To be fair, he does show this way once, but unfortunately, he keeps going back to the `for` loop throughout the rest of the text.

Not only is the second form shorter and clearer, but it also runs faster! Let's verify with the special jupyter command `%%timeit`. (It runs the cell many times to find the average run time.)

```python
%%timeit

xpoints = []
ypoints = []
for x in np.linspace(0, 10, 10000):
    xpoints.append(x)
    ypoints.append(np.sin(x))
```

```python
%%timeit

x = np.linspace(0, 10, 10000)
y = np.sin(x)
```