Fill in any place that says `# YOUR CODE HERE` or YOUR ANSWER HERE, as well as your name and collaborators below.
Grading for pre-lecture assignments is all or nothing. Partial credit is available for in-class assignments and checkpoints, but **only when code is commented**.

```python
NAME = ""
COLLABORATORS = ""
```

---

```python
import grading_helper as _test
```

<!-- grade_id: cell-76dabab5228e11a2, points: 0, grade: false, solution: true -->
```python
# Space for imports, utility functions, etc.

# YOUR CODE HERE
```

# A Random Point on the Surface of a Sphere

(Adapted from textbook exercise 10.12)

Suppose you wish to choose a random point on the surface of the Earth. That is, you want to choose a value of the latitude and longitude such that every point on the planet is equally likely to be chosen. In a physics context, this is equivalent to choosing a random vector direction in three-dimensional space (something that one has to do quite often in physics calculations). If we just chose uniformly distributed values of latitude and longitude, we'd end up with too many points near the poles (or $z$-axis), and not enough near the equator (the $xy$-plane).

Recall that in spherical coordinates $\theta,\phi$ the element of solid angle is $\sin \theta\,d\theta\,d\phi$, and the total solid angle in a whole sphere is $4\pi$. Hence the probability of our point falling in a particular element is
$$p(\theta,\phi)\,d\theta\,d\phi = \frac{\sin \theta \,d\theta\,d\phi}{4\pi}\,.$$
We can break this up into its $\theta$ part and its $\phi$ part thus:

$$p(\theta) = \frac{\sin \theta}{2}\,,$$

$$p(\phi) = \frac{1}{2\pi}\,.$$

1. Find formulas (by hand) for generating angles $\theta$ and $\phi$ drawn from the distributions $p(\theta)$ and $p(\phi)$. (The $\phi$ one is trivial, but the $\theta$ one is not.)
2. Write a program that generates a random $\theta$ and $\phi$ using the formulas you worked out.
3. Use your program to generate 10,000 such random points, convert the angles to $x,y,z$ coordinates assuming the radius of the sphere is 1, and save the results in three arrays named `x`, `y`, and `z`.

<!-- grade_id: cell-b22868219cde38a7, points: 0, grade: false, solution: true -->
```python
%%graded # 7 points

# YOUR CODE HERE
```

<!-- grade_id: cell-fb700eec3f6c1bb8, points: 7, grade: true, solution: false -->
```python
%%tests
_test.equal(len(x), 10000)
_test.equal(len(y), 10000)
_test.equal(len(z), 10000)

# check range
assert max(x) <= 1
assert min(x) >= -1
assert max(y) <= 1
assert min(y) >= -1
assert max(z) <= 1
assert min(z) >= -1

# check distirbution of value. The plot you make in the next part is a
# visual version of this test
import numpy as _np
_hist, _ = _np.histogram(x, bins=5)
_test.similar(_hist/2000, _np.ones(5), rtol = 0.1)
_hist, _ = _np.histogram(y, bins=5)
_test.similar(_hist/2000, _np.ones(5), rtol = 0.1)
_hist, _ = _np.histogram(z, bins=5)
_test.similar(_hist/2000, _np.ones(5), rtol = 0.1)
```

Plot the $(x, y)$ and $(x, z)$ coordinates of the points (as dots) to verify that they are properly distributed. Set `alpha=0.2` (transparency) in your plot to make the points more visible.

> Watch out for central spots or equitorial gaps in your images. Those are signs that the random numbers are not being drawn from the correct distribution.

<!-- grade_id: cell-f308fb92ff818f02, points: 2, grade: true, solution: true -->
```python
%%graded # 3 points

# YOUR CODE HERE
```

<!-- grade_id: cell-fa6f920261a68f5e, points: 1, grade: true, solution: false -->
```python
%%tests

_test.code_contains("alpha")
_test.plot_shown()
```