Grading for pre-lecture questions is all or nothing. Partial credit is available for in-class assignments and checkpoints. Partial credit will primarily be determined by the quality of the comments you make in your code. Tell me what you intend the code to do. If there are no comments, there is no partial credit!

```python
NAME = ""
COLLABORATORS = ""
```

---

# Learning Objectives

This lecture will show you how to:
1. Solve a nonlinear equation with the relaxation method
2. Apply the relaxation method to simultaneous equations
3. Implement the binary search algorithm
4. Implement the Newton-Raphson method

```python
# imports
import numpy as np
import matplotlib.pyplot as plt
from IPython import display
import time

from scipy import optimize # find equation roots and minima

import grading_helper as _test
```

# Relaxation Method

```python
%video X67U_8sxhGA
```

Summary:

- Rewrite your equation in the form $x=f(x)$. There may be more than one way to do so.
- The solution will converge if $|f^\prime(x_0)|<1$, where $x_0$ is the solution. Otherwise you need to write your equation differently.
- `scipy.optimize.fixed_point(f, guess)` implements this method with an enhancement called **over-relaxation**, that takes larger steps.

```python

```

## Your Turn

Use `optimize.fixed_point` to find the solution of 
$$x + \frac{1}{x} = 10$$
closest to $x=0$. Save the result in a variable named `x0`.

<!-- grade_id: cell-35ad44336baecb36, points: 0, grade: false, solution: true -->
```python
%%graded # 1 point

# YOUR CODE HERE
```

<!-- grade_id: cell-da8b5d94fab34508, points: 1, grade: true, solution: false -->
```python
%%tests

_test.code_contains("optimize.fixed_point")
_test.similar(x0 + 1/x0, 10)
assert x0 < 1 # check it isn't the other solution
```

# Relaxation Method With Simultaneous Equations

```python
%video wPlwl-H7e1k
```

Summary:

- Unlike the other methods we'll discuss, the relaxation method also works to solve simultaneous nonlinear equations with multiple variables.
- Recipe (2-D example)
    1. Write equations in the form
    $$x = f(x, y)$$
    $$y = g(x, y)$$
    2. Pick starting values for $x$ and $y$.
    3. Update $x$ and $y$ at the same time (use multiple assignment).
    4. Check **both** for convergence.
    5. Repeat from step 3 as needed.

```python

```

# Binary Search

```python
%video Crtjj96mMRs
```

Summary:

- In many cases, we know roughly where the solutions are and we just need to refine our estimate of them. We write the equation in the form $f(x) = 0$ and look for the roots:
    1. Find two points, $x_1$ and $x_2$ that you know are on either side of the value $x_0$ that satisfies $f(x_0) = 0$. It's important that $f(x_1) = -f(x_2)$.
    2. Let $x^\prime$ be the midpoint. Check the sign of $f(x^\prime)$ and then make a new interval by using $x^\prime$ in place of $x_1$ or $x_2$, depending on sign.
    3. Repeat step 2 until the interval is as small as your tolerance.
- This is a very fast method.
- Possible issues:
    - You need to know where the solution is well enough to bracket it.
    - If there are multiple solutions between $x_1$ and $x_2$, you'll only find one (although multiple solutions is an issue with all methods).
    - Fails for **double roots**. For example, $x^2 = 0$ has a double root at $x=0$, but binary search won't work.
- `scipy.optimize.bisect(f, x1, x2)` implements this method.

```python

```

## Your Turn

Use `optimize.bisect` to find the solution of 
$$x + \frac{1}{x} = 10$$
closest to $x=0$. Save the result in a variable named `x1`.

<!-- grade_id: cell-c85222addd195fd0, points: 0, grade: false, solution: true -->
```python
%%graded # 2 points

# YOUR CODE HERE
```

<!-- grade_id: cell-66bb15c80b0dc98c, points: 2, grade: true, solution: false -->
```python
%%tests

_test.code_contains("optimize.bisect")
_test.similar(x1 + 1/x1, 10)
assert x1 < 1 # check it isn't the other solution
```

# Newton-Raphson

```python
%video WbMUoSXYCh4
```

Summary:

- The first few terms of the Taylor expansion of $f(x)$ about the solution $x_0$ can be rearranged to
$$x_0 \approx x - \frac{f(x)}{f^\prime(x)}\,.$$
Since this relationship isn't exact, we can't actually find $x_0$ after just one iteration. However, we can use this result to hopefully improve our estimate. We just keep iterating
$$x = x - \frac{f(x)}{f^\prime(x)}\,.$$
- This method requires that we know $f^\prime(x)$, which isn't always going to be the case. It may occur to you to use a numerical derivative, and that's exactly what the **secant method** does.
- `scipy.optimize.newton(f, guess, fprime)` implements the Newton-Raphson method. Omitting `fprime` switches to using the secant method.

```python

```

## Your Turn

Use `optimize.newton` to find the solution of 
$$x + \frac{1}{x} = 10$$
closest to $x=0$. Save the result in a variable named `x2`.

<!-- grade_id: cell-19e236b9d80f706c, points: 0, grade: false, solution: true -->
```python
%%graded # 2 points

# YOUR CODE HERE
```

<!-- grade_id: cell-dd348913c3df22d2, points: 2, grade: true, solution: false -->
```python
%%tests

_test.code_contains("optimize.newton")
_test.similar(x2 + 1/x2, 10)
assert x2 < 1 # check it isn't the other solution
```

# Additional Resources

- Textbook section 6.3
- See exercise 6.11 for a walk-through of implementing over-relaxation.