Fill in any place that says `# YOUR CODE HERE` or YOUR ANSWER HERE, as well as your name and collaborators below.
Grading for pre-lecture assignments is all or nothing. Partial credit is available for in-class assignments and checkpoints, but **only when code is commented**.

```python
NAME = ""
COLLABORATORS = ""
```

---

# Learning Objectives

This lecture will show you how to:
1. Expand a function as a Fourier Series
2. Calculate the continuous Fourier transform of a function
3. Calculate the discrete Fourier transform of an array

```python
# imports
import numpy as np
import matplotlib.pyplot as plt
from scipy import integrate

import grading_helper as _test
```

# Fourier Series

```python
%video VxcsfveSEEg
```

Summary:

A periodic function $f(x)$ in the range $-L<x<L$ can be written in the form
$$f(x) = \sum_{k=-\infty}^\infty\gamma_k\exp\bigg(\mathrm{i}\frac{k\pi x}{L}\bigg)\,,$$
where the coefficients $\gamma_k$ are found using
$$\gamma_k = \frac{1}{2L}\int_{-L}^Lf(x)\exp\bigg(-\mathrm{i}\frac{k\pi x}{L}\bigg)\,dx\,.$$
Note that $\gamma_k$ is complex, even if $f(x)$ is real.

```python

```

# Continuous Fourier Transform

```python
%video dWvrXBI1zwg
```

Summary:

- Uses:
    - Break periodic function into component frequencies
    - Smooth and filter functions
    - Solve some differential equations

$$F(\omega) = \frac{1}{2\pi}\int_{-\infty}^\infty f(x)\,e^{-\mathrm{i} \omega x}\,dx\qquad\mbox{Fourier Transform of $f(x)$}\,,$$
and
$$f(x) = \int_{-\infty}^\infty F(\omega)\,e^{\mathrm{i} \omega x}\,d\omega\qquad\mbox{Inverse Fourier Transform of $F(\omega)$}\,.$$

```python

```

## Your Turn

Try to calculate the **inverse** Fourier transform of

$$F(\omega) = \frac{\omega\cos\omega - \sin\omega}{\pi\omega^2}\,\mathrm{i}\,,$$

which is the result found in the video. Use `quad` to evaluate the integral (remember that the integration limits are inifinite: $\pm$ `np.inf`). The result *should* be the original $f(x)$ from the video. Plot it to see that it doesn't work very well. Why do you think that is? These kinds of failures are why we typically use the **discrete Fourier transform** instead.

<!-- grade_id: cell-f379c1cb575251f8, points: 0, grade: false, solution: true -->
```python
%%graded # 2 points

# YOUR CODE HERE
```

<!-- grade_id: cell-59ddbd886466efc5, points: 2, grade: true, solution: false -->
```python
%%tests

_test.code_contains("quad", "cos", "sin", "exp")
_test.plot_shown()
```

# Discrete Fourier Transform

```python
%video eZrbVlIQRTI
```

Summary:

- Suppose we have an array $y$ rather than a function. Then to find the Fourier transform, we need to use a summation instead of an integration. We get exactly the same result if we evaluate the integral in the Fourier transform using the trapezoidal rule we find that $$c_k = N \gamma_k = \sum_{n=0}^{N-1}y_n\,\exp\bigg(-\mathrm{i}\frac{2\pi kn}{N}\bigg)\,,$$ where the $c_k$ are analgous to the $\gamma_k$ we found in the last lecture.
- 
The inverse transform is $$y_n = \frac{1}{N}\sum_{k=0}^{N-1}c_k\,\exp\bigg(\mathrm{i}\frac{2\pi kn}{N}\bigg)\,.$$
- Note that the $1/N$ term is in the inverse equation this time. The choice is actually arbitrary, but the version shown here is the normal convention.
- Since the DFT is generally complex, we typically plot its magnitude (using `abs`).

```python

```

## Your Turn

Calculate the **inverse** discrete Fourier transform of the array `c` below. Name your result `y`. 

```python
c = np.array([38.+0.j, 8.58694955-5.34676842j, 3.34332233+7.52563288j,
    -2.93027188+4.82461577j, -2.93027188-4.82461577j, 3.34332233-7.52563288j,
    8.58694955+5.34676842j])
```

<!-- grade_id: cell-35ad44336baecb37, points: 0, grade: false, solution: true -->
```python
%%graded # 3 points

# YOUR CODE HERE
```

<!-- grade_id: cell-da8b5d94fab34509, points: 3, grade: true, solution: false -->
```python
%%tests

_test.code_contains("exp", "pi")
_test.code_contains("fft", forbidden=True) # We'll use scipy next time. For now, code it yourself.
assert y.dtype == "complex" # the result is complex in general, even if it's real this time
_test.similar(y, np.array([8, 6, 7, 5, 3, 0, 9]), atol=1e-6)
```

# Additional Resources

- Textbook sections 7.1 and 7.2